Aeromonas hydrophila CobQ is a new type of NAD+- and Zn2+-independent protein lysine deacetylase

Author response:

The following is the authors’ response to the current reviews.

Public Reviews:

Reviewer #1 (Public review):

Summary:

This study by Wang et al. identifies a new type of deacetylase, CobQ, in Aeromonas hydrophila. Notably, the identification of this deacetylase reveals a lack of homology with eukaryotic counterparts, thus underscoring its unique evolutionary trajectory within the bacterial domain.

Strengths:

The manuscript convincingly illustrates CobQ's deacetylase activity through robust in vitro experiments, establishing its distinctiveness from known prokaryotic deacetylases. Additionally, the authors elucidate CobQ's potential cooperation with other deacetylases in vivo to regulate bacterial cellular processes. Furthermore, the study highlights CobQ's significance in the regulation of acetylation within prokaryotic cells.

Weaknesses:

The problem I raised has been well resolved. I have no further questions.

Thanks for your valuable comments very much.

Reviewer #2 (Public review):

In recent years, lots of researchers tried to explore the existence of new acetyltransferase and deacetylase by using specific antibody enrichment technologies and high resolution mass spectrometry. Here is an example for this effort. Yuqian Wang et al. studied a novel Zn2+- and NAD+-independent KDAC protein, AhCobQ, in Aeromonas hydrophila. They studied the biological function of AhCobQ by using biochemistry method and MS identification technology to confirm it. These results extended our understanding of the regulatory mechanism of bacterial lysine acetylation modifications. However, I find this conclusion is a little speculative, and unfortunately it also doesn't totally support the conclusion as the authors provided.

Major concerns:

- It is a little arbitrary to come to the title "Aeromonas hydrophila CobQ is a new type of NAD+- and Zn2+-independent protein lysine deacetylase in prokaryotes." It should be modified to delete the "in the prokaryotes" except that the authors get new more evidence in the other prokaryotes for the existence of the AhCobQ.

Thank you for your suggestion. However, I believe there has been some confusion regarding the title. In the revised manuscript we have already updated the title to: "Aeromonas hydrophila CobQ is a new type of NAD+- and Zn2+-independent protein lysine deacetylase."

This title does not include the phrase "in prokaryotes," as you mentioned. We kindly suggest verifying the version of the manuscript that was reviewed to ensure you are reviewing the most recent changes.

- I was confused about the arrangement of the supplementary results. Because there are no citations for Figures S9-S19.

Thank you for your feedback. It appears there may have been a misunderstanding, possibly due to reviewing an outdated version of the manuscript. In the revised manuscript we revised the supplementary figures and now have only 12 figures, all of which are correctly cited in the manuscript on pages 12 to 15. Below is a detailed list of the updated figure citations:

Figures S1: page 8, line 148;

Figures S2: page 9, line 168;

Figures S3 and S4: page 10, line 178;

Figures S5: page 10, line 186;

Figures S6: page 10, line 189;

Figures S7: page 12, line 221;

Figures S8-S10: page 13, line 245;

Figures S11: page 11, line 282;

Figures S12: page 15, line 286

- Same to the above, there are no data about Tables S1-S6.

Thank you for your attention to the supplementary materials. As with the figures, we have already uploaded the data for Tables S1-S6 in the revised manuscript on November 19, 2024, and properly cited Tables S1 – S6 in the manuscript. Below is the citation information:

Tables S1: page 10, line 194;

Tables S2: page 13, line 245;

Tables S3: page 21, line 438;

Tables S4: page 22, line 439;

Tables S5: page 22, line 445;

Tables S6: page 27, line 564.

Please note that Tables S3 – S4 include the chemical reagents, primers, and other experimental materials, which are not intended to be cited in the results section.)

- All the load control is not integrated. Please provide all of the load controls with whole PAGE gel or whole membrane western blot results. Without these whole results, it is not convincing to come the conclusion as the authors mentioned in the context.

Thank you for your comment. Please note that the full membrane western blot results were included in the revised manuscript. We hope this satisfies your request. If you need further clarification or additional data, please do not hesitate to let us know.

- Thoroughly review the materials & methods section. It is unclear to me what exactly the authors describe in the method. All the experimental designs and protocols should be described in detail, including growth conditions, assay conditions, and purification conditions, etc.

Thank you for your valuable suggestion. In response to your comment and previous feedback, we have alredy revised the Materials & Methods section thoroughly in the revised manuscript. The experimental details, including growth conditions, assay protocols, and purification procedures, are described in full on pages 22 to 30 of the revised manuscript.

- Include relevant information about the experiments performed in the figure legends, such as experimental conditions, replicates, etc. Often it is not clear what was done based on the figure legend description.

Thank you very much for your detailed feedback and suggestions. We have made sure to describe what each data point represents in the figure legends, as per the previous feedback. However, we would like to clarify that while we have provided detailed descriptions in the legends, the inclusion of every specific experimental condition in the figure legends could result in redundancy, as these details are already thoroughly outlined in the Materials & Methods section.

We hope this explanation addresses your concern.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

I have no further revision comments.

Thank you very much.

Reviewer #2 (Recommendations for the authors):

I carefully read the point-to-point response from the author. Although they listed lots of the reasons for the ugly results, it still can not persuade me to accept their conclusions. While, as I know, it is impossible to reject their work in eLife as it was sent out for peer-review. I also can't accuse them of being wrong, but I have my opinion on this point. That is not the results, but the attitude.

Thank you for your feedback. However, I must express some concerns regarding the nature of your comments. Based on the issues you've raised, it seems that you may have reviewed an outdated version of the manuscript. In the updated revision we addressed all the points you've raised, including the figure and table citations, experimental methods, and data integration.

We understand that differing opinions are part of the peer-review process, but we respectfully believe that your conclusion regarding our attitude is based on a misunderstanding, possibly caused by reviewing an incorrect version of the manuscript. We have always strived to approach this manuscript with utmost professionalism and have diligently responded to each of your concerns.

We sincerely suggest reviewing the latest version of our manuscript, and we welcome any further constructive feedback. We hope this clarifies any misunderstandings and look forward to your continued support.

Thank you for your time and thoughtful consideration.

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1 (Public review):

Summary:

This study by Wang et al. identifies a new type of deacetylase, CobQ, in Aeromonas hydrophila. Notably, the identification of this deacetylase reveals a lack of homology with eukaryotic counterparts, thus underscoring its unique evolutionary trajectory within the bacterial domain.

Strengths:

The manuscript convincingly illustrates CobQ's deacetylase activity through robust in vitro experiments, establishing its distinctiveness from known prokaryotic deacetylases. Additionally, the authors elucidate CobQ's potential cooperation with other deacetylases in vivo to regulate bacterial cellular processes. Furthermore, the study highlights CobQ's significance in the regulation of acetylation within prokaryotic cells.

Weaknesses:

The problem I raised has been well resolved. I have no further questions.

Reviewer #2 (Public review):

In recent years, lots of researchers tried to explore the existence of new acetyltransferase and deacetylase by using specific antibody enrichment technologies and high resolution mass spectrometry. Here is an example for this effort. Yuqian Wang et al. studied a novel Zn2+- and NAD+-independent KDAC protein, AhCobQ, in Aeromonas hydrophila. They studied the biological function of AhCobQ by using biochemistry method and MS identification technology to confirm it. These results extended our understanding of the regulatory mechanism of bacterial lysine acetylation modifications. However, I find this conclusion is a little speculative, and unfortunately, it also doesn't totally support the conclusion as the authors provided.

Reviewer #3 (Public review):

Summary:

This study reports on a novel NAD+ and Zn2+-independent protein lysine deacetylase (KDAC) in Aeromonas hydrophila, termed as AhCobQ (AHA_1389). This protein is annotated as a CobQ/CobB/MinD/ParA family protein and does not show similarity with known NAD+-dependent or Zn2+-dependent KDACs. The authors showed that AhCobQ has NAD+ and Zn2+-independent deacetylase activity with acetylated BSA by western blot and MS analyses. They also provided evidence that the 195-245 aa region of AhCobQ is responsible for the deacetylase activity, which is conserved in some marine prokaryotes and has no similarity with eukaryotic proteins. They identified target proteins of AhCobQ deacetylase by proteomic analysis and verified the deacetylase activity using site-specific Kac proteins. Finally, they showed that AhCobQ activates isocitrate dehydrogenase by deacetylation at K388.

Strengths:

The finding of a new type of KDAC has a valuable impact on the field of protein acetylation. The characters (NAD+ and Zn2+-independent deacetylase activity in an unknown domain) shown in this study are very unexpected.

Weaknesses:

(1) The characters (NAD+ and Zn2+-independent deacetylase activity in an unknown domain) shown in this study are very unexpected. To convince readers, MSMS data must be necessary to accurately detect (de)acetylation at the target site in the deacetylase activity assay. The authors showed the MSMS data in assays with acetylated BSA, but other assays only rely on western blot.

(2) They prepared site-specific Kac proteins and used them in deacetylase activity assays. Incorporation of acetyllysine at the target site should be confirmed by MSMS and shown as supplementary data.

(3) The authors imply that the 195-245 aa region of AhCobQ may represent a new domain responsible for deacetylase activity. The feature of the region would be of interest but is not sufficiently described in Figure 5. The amino acid sequence alignments with representative proteins with conserved residues would be informative. It would be also informative if the modeled structure predicted by AlphaFold is shown and the structural similarity with known deacetylases is discussed.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

The problem I raised has been well resolved. I have no further questions.

Reviewer #2 (Recommendations for the authors):

Questions to response of"-The load control is not all integrated. All of the load controls with whole PAGE gel or whole membrane western blot results should be provided. Without these whole results, it is not convincing to come to the conclusion that the authors have."

Just as the Authors answered. The Coomassie Blue R-350 staining outcomes from the PVDF membranes. That is a good control for the experiment. However, I still have several questions about it:

(1) The first is the quality of these Western blot. Why all the bands of these Western blot is so ugly? To tell the truth, it is very difficult to come to a conclusion from these poor western blots.

We appreciate your feedback regarding the quality of the Western blots presented in Figure 7. We believe the “ugly bands” you referred to reflect our results validating the functions of CobQ through the use of recombinant site-specific Kac protein substrates.

In our study, we meticulously engineered these recombinant site-specific Kac proteins using a two-plasmid system, based on foundational research published in Nature Chemical Biology (2017, 13(12): 1253-1260), which introduced the genetic encoding of Nε-acetyllysine into recombinant proteins. However, we faced a common challenge: protein truncation due to premature translation termination at the reassigned codon. This issue not only hampers protein yields, as discussed in ChemBioChem (2017, 18(20): 1973-1983), but also contributes to the suboptimal appearance of the Western blot results.

Despite conducting at least two independent repetitions for the Western blot analysis of the site-specific Kac proteins, which yielded consistent results, we recognize that the overall quality remains less than ideal. This variability is inherently related to the characteristics of the target proteins. Nevertheless, the primary aim of our manuscript is to validate the novel deacetylase activity of CobQ. We have provided multiple lines of evidence, including mass spectrometry (MS/MS) and Western blot analyses, to substantiate this claim. In response to your comments, we have decided to remove the ambiguous Western blot results from Figure 7, retaining only four figures that demonstrate significant differences across at least two independent replicates (Author response images 1-5). Additionally, we have included four biological replicates of the Western blot results for ICD Kac388 + CobQ in the supplementary materials (Author response image 5) to further validate the deacetylase function of CobQ.

Western blot validation of the Kac26 AcrA-2 protein substrates regulated by the three KDACs in two biological replicates.

Western blot validation of the Kac48 Sun protein substrates regulated by the three KDACs in two biological replicates.

Western blot validation of the Kac103 Sun protein substrates regulated by the three KDACs in two biological replicates.

Western blot validation of the Kac195 Eno protein substrates regulated by the three KDACs in three biological replicates.

Western blot validation of the Kac388 ICD protein substrates regulated by AhCobQ in this study.

Each sample was independently repeated at least three time.

(2) The second is why some of the results are not from the same PVDF by comparing the Coommassie staining with the WB results just as authors responded. For example, the HrpA-K816 (ac), Eno-K195 (ac), ArcA-2-K26 (ac), ArcA-2-K26 (ac), IscS-K93(ac), A0KJ75-K81(ac), GyrB-K331(ac), GyrB-K449(ac), FtsA-K320(ac), FtsA-K409(ac), RecA-K279(ac), and the RecA-K306(ac). All of them are clearly not from the same staining results of PVDF membrane but from a new PVDF membrane.

We assure you that the R-350 stained PVDF membranes originate from the same Western blot membranes. However, we acknowledge that visual discrepancies may arise due to differences in imaging techniques. The Western blot results were scanned using a ChemiDoc MP (Bio-Rad, Hercules, CA, USA), while the Coomassie R-350 stained PVDF membranes were captured using a standard camera. These differences can create a misleading appearance, making it seem as though they come from different membranes.

It is also important to note that the intensity of the protein marker cannot be directly compared between the two imaging methods. As illustrated in Author response image 6, the protein marker at 70 kDa is clearly detectable in the Coomassie R-350 image, whereas it may not be as apparent in the Western blot result due to inherent differences in detection sensitivity.

The comparison of Western blotting and R-350 strained results of same protein marker in the same PVDF membrane.

The protein marker located at 70 kDa can be detected easily in Coomassie R-350, while is difficult to display in WB result.

Additionally, we have removed some of the so-called "ugly" Western blot results in the updated manuscript and provided the original full film of the relevant images as an attachment. This documentation demonstrates that all the data you referenced originate from the same film, as shown in Figures 1-5.

(3) The third is why there is no replication for all these WB results. We should draw a conclusion with serious attitude, but not from the only one repeat, even say nothing about the poor results.

Thank you for your valuable suggestion. In the second version of the manuscript, we have included the original full film of the relevant images. While we previously explained the reasons behind the "ugly" Western blot results, we have decided to remove some, or even all, of these results from Figure 7 in the updated version. The related images will be updated in the supplementary materials (Figures 1-5 in responding letter and Figure 7 in the revised manuscript).

Furthermore, we have provided a more detailed discussion regarding the poor results in the updated manuscript to ensure clarity and transparency. We appreciate your understanding and hope these changes meet your expectations.

Questions to response of " L174-187, L795 (Please show the whole membrane (or PAGE gel) of the loading control of CobB, and CobQ, except for the Kac-BSA)".

(1) As we all educated that there is no control, and no biology. Where is the band of CobQ? Why do not stain the same PVDF membrane with R-350 staining but with a new membrane?

Thank you for your insightful feedback. As noted in our previous response, the absence of visible bands for AhCobQ and AhCobB on the Coomassie R-350 stained PVDF membrane is primarily due to the low loading amounts and protein loss during the Western blotting process.

To reinforce our findings, we repeated the analysis of the protein samples via SDS-PAGE, using the same loading quantity as in the previous Western blot shown in Figure 2 of the manuscript. As illustrated in Author response image 7, the bands for CobB and CobQ are discernible, albeit with significantly lower intensities compared to the Kac-BSA bands. Upon examining the full Coomassie R-350 stained PVDF membranes provided in Supplementary Material 1, we observe that the CobB and CobQ bands are not easily visible. This aligns with your observations and can be attributed to potential protein loss during the transfer from SDS-PAGE to the PVDF membrane.

The SDS-PAGE gel displayed the loading amounts of Kac-BSA and CobB/CobQ.

To enhance the visibility of the CobQ/CobB bands, we increased the loading of CobQ/CobB in a new Western blot experiment, using 2 µg of Kac-BSA in combination with 0.8 µg of CobQ/CobB. As shown in Figure 8, while the increasing amounts of Kac-BSA resulted in a more blurred signal, the bands for the recombinant CobQ and CobB proteins were clearly detectable. This indicates that both proteins were indeed involved in the in vitro protein deacetylation assay.

Western blot verified the deacetylase activity assay of AhCobQ and AhCobB on Kac-BSA.

Furthermore, we conducted a mass spectrometry analysis comparing Kac-BSA and Kac-BSA incubated with CobQ, as well as BSA without acetylation, against the A. hydrophila database with a cut-off of unique matched peptides >1. It is challenging to completely avoid contaminant detection during protein purification, especially when using high-resolution mass spectrometry. Our findings revealed that CobQ has the highest number of unique matched peptides (Author response table 1), while contaminants such as AHA_3036, AHA_0497, AHA_1279, and valS could be excluded, as they were present in Kac-BSA or BSA samples. Additionally, Tuf1, RplQ, GroEL, RpsF, RpsU, RpsB, RpsO, and RpsJ are known ribosomal subunits or chaperonins that are abundantly expressed in cells and may interact with various proteins, leading to contaminant detection.

Although AceE, a pyruvate dehydrogenase E1 component, theoretically possesses deacetylase activity, this possibility is low. First, in the SDS-PAGE gel of the purified recombinant protein, CobQ is the major band, with other proteins present at very low levels (less than 1/10 of CobQ). This suggests that significant deacetylation by contaminants is unlikely. Second, we purified His-tagged AhCobQ and GST-fused AhCobQ separately and tested their deacetylase activities. As shown in Figure S4 of the updated manuscript, both purified AhCobQ proteins exhibited deacetylase activity, while the negative control (purified GST protein only) did not, further supporting our conclusion that enzyme activity is not attributable to contaminating proteins (Figure S5).

(2) Without the CobB and CobQ bands, it is impossible to say the function of CobQ is a new deacetylase. To avoid this confusion, it is easy to run a new gel and stain it with anti-His antibody to show these deacetylases.

Thank you very much for your suggestion. We have performed the experiment in the comment (1) as your suggestion.

(3) The explanation about the CobB/CobQ bands are not visible is not acceptable. Because the molecular weight of the CobB and CobQ is smaller than that of BSA, it is impossible that these bands will be loss during membrane transfer.

Thank you for your valuable feedback. I completely agree that the loss of CobB and CobQ proteins during membrane transfer is unlikely due to their smaller molecular weight compared to BSA. As shown in Figure 7, the bands for CobB and CobQ are detectable in the SDS-PAGE gel but not visible on the Coomassie R-350 stained PVDF membrane.

Several factors could contribute to this issue. One possibility is that the detection sensitivity of Coomassie R-350 may be lower than that of Coomassie R-250 used in the gel. Additionally, the Western blot results using an anti-His antibody further indicate low loading amounts of CobB and CobQ proteins on the PVDF membrane (Figure 8). This suggests that the observed low levels may indeed be due to protein loss during the membrane transfer process, despite their relatively small size.

Reviewer #3 (Recommendations for the authors):

(1) I found Tables S1 and S2 in the revised manuscript. It is strange to me that the intensity of Kac-BSA+CobQ is zero, completely nothing. Typically, a portion of the acetylated peptide remains after the deacetylation reaction.

Thank you for your observation. When we report an intensity of zero, it does not imply a complete absence of signal; rather, it indicates that the signal for the target peptide is below the detectable threshold. This is likely due to the minimum cut-off setting in the MaxQuant (MQ) software, which is determined by parameters like "peptide_mass_tolerance" (as discussed in MQ user groups online, though it may not be explicitly listed in the parameters file).

In our study, we performed a deacetylase assay that demonstrated CobQ's rapid activity; for instance, it can deacetylate ICD-K388ac within just four minutes. This leads me to hypothesize that the CobQ + Kac-BSA sample may have undergone near-complete enzymatic hydrolysis during the reaction.

Furthermore, Table S1 in manuscript presents only a selection of the mass spectrometry results to illustrate CobQ's activity. In addition to the 15 acetylated peptides shown, there are many more (27 peptides) that exhibit significantly reduced acetylation levels without reaching zero intensity. The overall acetylation level of BSA peptides incubated with CobQ is calculated to be only 0.13 times that of Kac-BSA (Diagnostic peak: yes, peptide score: >100, Localization probability: >0.95) (Author response image 9).

Based on these findings, we believe our mass spectrometry results are reliable and effectively support our conclusions. Thank you for your understanding.

The intensities of all Kac peptides of Kac-BSA with or without AhCobQ incubation in LC MS/MS.

(2) It would be better to provide the information about ArcA and ArcA-2 as mentioned in the authors' response. It would be helpful for readers to understand that they are different proteins.

Thank you for your suggestion. In the A. hydrophila ATCC 7966 dataset, there are indeed two distinct proteins referred to as ArcA: ArcA-1, which functions as an aerobic respiration control protein, and ArcA-2, which acts as an arginine deiminase. Importantly, these two proteins do not share any sequence homology; they are only similarly named due to their acronyms. While we believe this distinction does not require extensive explanation in the current study, we appreciate your input. Additionally, in response to Reviewer 2’s feedback, we have decided to remove the Western blot result for ArcA-2 due to its poor quality in the updated manuscript.

(3) Line 409-416. Despite my comment, the citation of related papers on ICD acetylation in E. coli is still missing.

Thank you for your suggestion. It has been added and highlighted in red. (Venkat S, et al, 2018, 430(13): 1901-1911)

(4) The image resolution of Figure 3C and 3D is still bad. I could not evaluate that Kac was exactly incorporated at the target site.

Thank you for your feedback regarding the image resolution of Figures 3C and 3D. We have now displayed these figures with improved clarity, as you suggested.

To further validate the reliability of our MS2 data, we employed Proteome Discoverer 2.4 (Thermo) to analyze the raw data and provide theoretical mass information. As shown in Author response images 10-13, the MS2 spectra and fragment match lists for both unmodified and acetylated peptides offer additional confirmation of the reliability of our mass spectrometry results.

MS2 spectrum of unmodified peptide using PD v2.

4 software.

The theoretical mass of unmodified peptide by PD 2.

4

MS2 spectrum of acetylated peptide using PD v2.

4 software.

The theoretical mass of acetylated peptide by PD 2.

4.

(5) Again, in Figure 8D, it should be shown the significance between ICD-Kac388 and ICD-Kac388+AhCobB to support the authors' conclusion that AhCobQ activates ICD by deacetylation at K388.

Thanks for your suggestion, we have updated the figure in Figure 8D in updated manuscript.

(6) It was nice that the authors presented the mass spectrum data of ICD-K388 acetylation (Figure 2 in responding letter). However, the data did not convince me that K388 is acetylated. In the figure, two b-ion peaks are detected, 285.1557 and 386.2034, which may correspond to NK (theoretical mass, 260.15) and NKT (theoretical mass, 361.20) peptides, respectively. If K388 is acetylated, an increase in the mass of 42 should be observed, but the difference between the detected and theoretical mass is 25. I also could not understand what the peak of 126.0913 mass is, indicated with acK* in red.

Thank you for your detailed observation. The data presented in the MS2 spectrum for ICD-K388 acetylation in Figure 2 of the previous response letter were generated using Proteome Discoverer 2.4 (PD, Thermo) to ensure accurate mass calculations. Similar to the results from MaxQuant, ICD-K388 was identified again (Author response image 14).

Regarding the b-ion peaks you mentioned, the values 285.1557 and 386.2034 correspond to NKac and NKacT peptides, respectively. The theoretical masses for these peptides are as follows: NKac (285.15 = 115.05020 + 128.095 + 42.01) and NKacT (386.20 = NKac + 101.04768). The differences between the theoretical and detected masses for the relevant b-ions (b2*-NK, b52+-NH3, and b3) are minimal, at 0.00 Da and 2.1 ppm, respectively, which is consistent with the incorporation of an NH3 group (Author response image 15).

The MS2 of ICD-K388 peptide by PD 2.

4.

The theoretical mass of ICD-K388 peptide by PD 2.

4.

The peak at 126.0913 m/z, indicated as acK*, represents immonium ions of ε-N-acetyllysine, which are generated during the fragmentation of acetyllysine. This diagnostic ion is widely recognized as a marker for identifying acetylated peptides (Nakayasu, et al,. A method to determine lysine acetylation stoichiometries. International journal of proteomics. 2014;2014(1):730725; Trelle et al., Utility of immonium ions for assignment of ε-N-acetyllysine-containing peptides by tandem mass spectrometry. Analytical chemistry. 2008;80(9):3422-30). Additionally, it is a default parameter in MaxQuant for identifying Kac peptides (Author response image 16).

Based on these findings, we believe the evidence supporting ICD-K388 acetylation is robust.

The default parameter in Kac peptide identification in Maxquant v1.

6 software

(7) As mentioned by other reviewers, some of the figures and tables are incomplete. Some panels (ex. Figure 7C and 7D) and explanations (ex. What are lanes 1, 2, and 3 in Figure S3) are still missing.

Thank you for your suggestion. It has been added.

Changes from the earlier version to the revised version

Abstract

Abstract

Protein N−-lysine acetylation (Kac) modifications play crucial roles in diverse physiological and pathological functions in cells. In prokaryotic cells, there are only two types of lysine deacetylases (KDACs) that are Zn2+- or NAD+-dependent. In this study, we reported a protein, AhCobQ, in Aeromonas hydrophila ATCC 7966 that presents NAD+- and Zn2+-independent KDAC activity. Furthermore, its KDAC activity is located in an unidentified domain (from 195–245 aa). Interestingly, AhCobQ has no homology with current known KDACs, and no homologous protein was found in eukaryotic cells. A protein substrate analysis showed that AhCobQ has specific protein substrates in common with other known KDACs, indicating that these KDACs can dynamically co-regulate the states of Kac proteins. Microbiological methods employed in this study affirmed AhCobQ’s positive regulation of isocitrate dehydrogenase (ICD) enzymatic activity at the K388 site, implicating AhCobQ in the modulation of bacterial enzymatic activities. In summary, our findings present compelling evidence that AhCobQ represents a distinctive type of KDAC with significant roles in bacterial biological functions.

Introduction

In recent years, due to the development of specific antibody enrichment technologies and high-resolution mass spectrometry, the research on protein post-translational modifications (PTMs) is increasing dramatically, and it provides a new platform for the study of bacterial function and regulation. PTMs refer to small molecular moieties that dynamically affect the structure and function of proteins through chemical modification of specific amino acid residues, so as to regulate many important biological processes, including cell metabolism, DNA transcription and replication, host infection, and stress resistance, and therefore play an important role in almost all physiological functions (Carabetta et al, 2007). Among them, the N−-lysine acetylation (Kac) modification on protein is widely found in bacteria and participates in almost all key biological processes in bacterial life activities and is attracting the attention of microbiologists (Yu et al, 2023).

Although there are many reports on the identification of acetylation modification patterns of bacterial protein lysine residues in recent years, two fundamental questions have not been solved: how do these acetylation modifications arise (acetylation modifiers; writers) and how are they removed (deacetylation modifiers; erasers) (Wang et al, 2020). On the one hand, acetyl groups can be transferred to the target protein by acetyltransferases or non-enzymatic (chemical) mechanisms (VanDrisse & Escalante-Semerena, 2019); on the other hand, proteins need to remove acetyl groups from modified proteins to reversibly regulate the function of target proteins. At present, the research on the functions and types of lysine deacetylases (KDACs) has been well documented in eukaryotes, and most of them are histone deacetylases (HDACs). There are currently at least 18 types of HDACs in four classes (I–IV) in mammals and four plant-specific histone deacetylase 2 (HD2) members in plant (Yoshida et al, 2017; Chen et al, 2020). Of these, HDAC types I, II, and IV are Zn2+-dependent, while type III deacetylases are NAD+-dependent sirtuins. Moreover, sirtuins can be divided into seven types (SIRT 1–7), which are distributed in different compartments of eukaryotic cells and perform important biological functions (Narita et al, 2019). Beside these, ABHD14B, a novel KDAC, was reported to deacetylate non-histone substrates and influences glucose metabolism in mammals (Rajendran et al, 2022; Rajendran et al, 2019).

The currently known KDACs in bacteria include Zn2 +-dependent RpLdaA or Kdac1 (homologous to eukaryotic type II KDACs), Zn2+-dependent AcuC (homologous to type I), and NAD+-dependent sirtuin (homologous to type III) family proteins (Watson et al, 2023). Sirtuins, represented by Escherichia coli CobB, are the most widely distributed KDACs in many bacterial species and are involved in the regulation of a variety of important physiological functions (Burckhardt et al, 2019; Liu et al, 2018; Wang et al, 2020). However, since the first report of the deacetylase function of CobB in Salmonella typhimurium in 1998, only two types (three classes) of KDACs have been confirmed in bacteria (Tsang & Escalante-Semerena, 1998; Macek et al, 2019). The reasons that hinder new KDAC discovery in bacteria may be the following: (1) due to the in-depth study of KDACs in eukaryotes, almost all the current research on bacterial deacetylases are based on the characteristics of their homologous proteins in eukaryotes, but less attention is paid to bacteria-specific KDACs; and (2) the deacetylation process between protein substrates and KDACs may be a transient and reversible weak interaction. Moreover, there are thousands of lysine acetylation modification sites in cellular proteins (Chen et al, 2022); so, it is difficult to speculate the characteristics of their specific binding deacetylase, which makes the efficiency of enriching specific KDACs by common protein-protein interaction methods, such as co-immunoprecipitation and pull-down technologies, challenging. As a result, the identification and mechanism of new deacetylases in prokaryotic cells are lagging. Therefore, the screening and identification of new deacetylases, especially for new types of KDACs that may not depend on NAD+ or Zn2+, have been a bottleneck limiting the research on the regulatory mechanisms of bacterial acetylation/deacetylation modifications.

In this study, we fortuitously identified the protein AhCobQ in Aeromonas hydrophila ATCC 7966 (UniProt ID A0KI27, gene ID AHA_1389) that presents NAD+- and Zn2+-independent protein lysine deacetylase activity. Interestingly, although this protein is described as a CobQ/CobB/MinD/ParA family protein in the UniProt database, it is homologous with cobyric acid synthase in other bacterial species but shares no homology with CobB sirtuin family proteins. In addition to deacetylating lysine acetylated sites common to AhCobB (UniProt ID A0KKN9, gene ID AHA_2317) sirtuin and another deacetylase AhAcuC (UniProt ID A0KLZ2, gene ID AHA_2788) in A. hydrophila, we also confirmed that AhCobQ deacetylated specific lysine acetylated sites as did other KDACs, even in the same protein. In general, our results demonstrated a new type of bacterial KDAC, which shares no homology with currently known lysine deacetylases, including those in eukaryotic cells, and extends our understanding of the complicated regulatory mechanism of bacterial lysine acetylation modifications.

Results

The physiological phenotypes of AhCobQ are different from those of AhCobB

Earlier version: This project proceeded from the misunderstanding that AhCobQ and AhCobB may be homologous proteins because AhCobQ is annotated as a CobQ/CobB/MinD/ParA family protein in the UniProt database. Therefore, we constructed ahcobB and ahcobQ mutant strains (Figure 1A) and compared their general physiological phenotypes. When compared to the wild-type strain (WT), the deletion of ahcobB did not affect thehemolytic or swarming ability of A. hydrophila. Notably, ΔahcobQ did not affect the hemolytic ability either but significantly increased the swarming ability (Figure 1B, C and Figure S1). Most interestingly, ΔahcobB sharply increased, while ΔahcobQ significantly decreased, the bacterial biofilm formation ability (Figure 1D). These results indicated that CobB and CobQ in A. hydrophila may play different roles in biological functions.

Revised version: This project proceeded from the misunderstanding that AhCobQ and AhCobB may be homologous proteins because AhCobQ is annotated as a CobQ/CobB/MinD/ParA family protein in the UniProt database. Therefore, we constructed ahcobB and ahcobQ mutant strains (Figure 1A) and compared their general physiological phenotypes. When compared to the wild-type strain (WT), the deletion of ahcobB did not affect the hemolytic or swarming ability of A. hydrophila. Notably, ΔahcobQ did not affect the hemolytic ability either but significantly increased the swarming ability (Figure 1B, C and Figure S1). Most interestingly, ΔahcobB sharply increased, while ΔahcobQ significantly decreased, the bacterial biofilm formation ability (Figure 1D). These results indicated that CobB and CobQ in A. hydrophila may play different roles in biological functions.

Added in the revised version: Phenotypes of ΔahcobB and ΔahcobQ strains

Earlier version: Phenotypes of ΔahcobB and ΔahcobQ strains. (A) Construction of ahcobB- and ahcobQ-defective strains. M: Marker; Lanes 1 and 3: PCR products of WT using the P7/P8 primer pairs of ahcobB and ahcobQ, respectively; Lanes 2 and 6: PCR products of ΔahcobB and ΔahcobQ, respectively, using P7/P8 primer pairs; Lanes 5 and 7: PCR products of WT using P5/P6 primer pairs of ahcobB and ahcobQ, respectively; Lanes 4 and 8: PCR products of ΔahcobB and ΔahcobQ, respectively, using the P5/P6 primer pairs. (B) Bacterial migration ability; (C) hemolytic activity; (D) histogram of the biofilm formation ability (OD 595 nm). **P &lt; 0.005.

Revised version: (A) Construction of ahcobB- and ahcobQ-defective strains. M: Marker; Lanes 1 and 3: PCR products of WT using the P7/P8 primer pairs of ahcobB and ahcobQ, respectively; Lanes 2 and 6: PCR products of ΔahcobB and ΔahcobQ, respectively, using P7/P8 primer pairs; Lanes 5 and 7: PCR products of WT using P5/P6 primer pairs of ahcobB and ahcobQ, respectively; Lanes 4 and 8: PCR products of ΔahcobB and ΔahcobQ, respectively, using the P5/P6 primer pairs. (B) Bacterial migration ability; (C) hemolytic activity; (D) histogram of the biofilm formation ability (OD 595 nm). **P < 0.005.

AhCobQ and AhCobB are not homologous proteins

We then further investigated the evolutionary relationship between AhCobQ and AhCobB. To our surprise, we found that the proteins belong to different protein families. AhCobB (UniProt ID A0KKN9) belongs to a well-known NAD+-dependent protein deacylase sirtuin family and shares a highly homologous SIR2 domain with CobB sirtuins in many bacterial species, such as E. coli (NPD_ECOLI), S. typhimurium (NPD_SALTY), Pseudomonas nitroreducens (A0A246F352_9PSED), and even seven types of sirtuins in humans (Figure 2A). However, AhCobQ (UniProt ID A0KI27) shares an AAA_31 ATPase domain with several homologous proteins, such as cobalamin biosynthesis protein CobQ in P. nitroreducens (A0A246FC44_9PSED), probable plasmid partitioning protein in P. aeruginosa (G3XCW7_PSEAE), and chromosome partitioning ATPase in S. typhimurium (A0A0F6B556_SALT1). Moreover, AhCobQ also shares homology with Soj protein (A0KQY7_AERHH) and cell division inhibitor MinD (A0KK56_AERHH) in A. hydrophila as well (Figure 2B and Figure S2). Interestingly, this protein was not found to have homologous proteins in humans or mice species, indicating the significant evolutionary difference between AhCobQ and sirtuins.

Homology comparison and deacetylase activity assay of AhCobQ and AhCobB

(A) and (B) Homologous and conserved domain analysis of AhCobQ and AhCobB family proteins. (C) and (D) Dot blot and western blot verified the whole cell protein Kac level among WT, ΔahcobB, and ΔahcobQ strains. (E–H) Effect of NAD+, NAM, Zn2+, and ATP on KDAC enzymatic activity of AhCobQ.

AhCobQ is an NAD- and Zn2+-independent lysine deacetylase

Dot blot and western blot experiments showed that the deletion of ahcobQ or ahcobB both significantly increased the whole protein lysine acetylation levels in A. hydrophila cells (Figure 2C and D), which aroused our interest as to whether AhCobQ is a lysine deacetylase or not. To test this possibility, we purified AhCobQ, AhAcuC, and AhCobB proteins and incubated them in vitro with lysine acetylated-bovine serum albumin (Kac-BSA) as a protein substrate with and without NAD+ or Zn2+ treatment (Figure S3 and S4). As a positive control, AhCobB and AhAcuC significantly deacetylated Kac-BSA in the presence of NAD+ or Zn2+ buffer, which met our expectation. However, AhCobQ clearly deacetylated Kac-BSA regardless of whether NAD+ or Zn2+ were present (Figure 2E and G). Moreover, the KDAC activity of AhCobB was inhibited when the sirtuin inhibitor nicotinamide (NAM) was added, but this did not affect the deacetylase activity of AhCobQ (Figure 2F). Furthermore, the introduction of the metal ion chelator, EDTA, had no discernible effect on the deacetylase activity of AhCobQ (Figure S5). These results strongly indicate that AhCobQ is an NAD+- and Zn2+-independent protein deacetylase. Additionally, since AhCobQ has an ATPase domain at the N-terminal, we also tested its deacetylase activity with or without ATP (Figure 2H and Figure S6). The results demonstrated that the KDAC activity of AhCobQ was ATP-independent, suggesting its enzymatic activity was downstream of the ATPase domain.

We further digested Kac-BSA to peptides after AhCobQ incubation and then subjected the peptides to LC MS/MS to identify their acetylated sites (Table S1). The results showed that the MS intensities of at least 15 Kac peptides in BSA were significantly decreased after AhCobQ treatment. To further validate the KDAC characteristics of AhCobQ and exclude the possibility that AhCobQ may hydrolyze Kac peptides or switch the lysine acetylation to other unknown PTMs, we selected and synthetized an acetylated peptide sequence (YNGVFQECCQAEDKacGACLLPK) with its unacetylated peptide as negative control and subjected them to additional deacetylation assays. The synthetized acetylated (STD*Kac) and deacetylated (STD) peptides were incubated with AhCobB and AhCobQ and then identified by LC MS/MS. As shown in Figure 3, the masses of the STD*Kac and STD peptides were 786.35538 (3+) and 772.35175 (3+), respectively. When the STD*Kac peptide was incubated with AhCobB or AhCobQ, the corresponding retention time in the C18 chromatography and the MS1 peak of the deacetylated peptide were found in both treatments (Figure 3A and B). Moreover, the deacetylated and acetylated peptides were further validated in MS2 by LC MS/MS (Figure 3C and D). These results indicate that AhCobQ is not a hydrolase that hydrolyzes target proteins but likely a protein deacetylase similar to AhCobB.

Added in the revised version: Validation of the KDAC activity of AhCobQ by LC-MS/MS

Earlier version: Validation of the KDAC activity of AhCobQ by LC-MS/MS. (A) Retention time of synthetic peptides YNGVFQECCQAEDKGACLLPK (STD) and YNGVFQECCQAEDKacGACL-L-PK (STD*Kac) with or without AhCobB and AhCoQ treatment; (B) MS1 spectrum of synthetic peptides STD and STD*Kac with or without AhCobB and AhCobQ treatment; (C) and (D) MS2 spectrum of acetylated and unmodified peptides.

Revised version: (A) Retention time of synthetic peptides YNGVFQECCQAEDKGACLLPK (STD) and YNGVFQECCQAEDKacGACL-L-PK (STD*Kac) with or without AhCobB and AhCoQ treatment; (B) MS1 spectrum of synthetic peptides STD and STD*Kac with or without AhCobB and AhCobQ treatment; (C) and (D) MS2 spectrum of acetylated and unmodified peptides.

The deacetylase activity of AhCobQ is in an unidentified domain

Because AhCobQ has a conservative N-terminal AAA family ATPase domain (1–179 aa) and its KDAC activity was not affected by ATP, we then asked whether this domain has deacetylase activity. We divided and purified recombinant AhCobQ protein into two truncated protein fragments of 1–179 aa and 179–264 aa and then tested their deacetylase activities. To our surprise, the 179-264 aa protein fragment but not the 1-179 aa fragment showed NAD+- and Zn2+-independent deacetylase activity (Figure 4A). To better understand the deacetylase domain of AhCobQ, a total of 12 truncated protein fragments were constructed and purified, and their deacetylase activities were evaluated (Figures S7 and 4A). As shown in Figure 4, the 1–179 aa, 200–255 aa, and 189–240 aa shortened proteins disrupted the deacetylase activity of AhCobQ to some extent, whereas the 179–264 aa, 189–264 aa, 199–264 aa, 195–264 aa, 195–255 aa, 189–245 aa, 189–250 aa, 189–255 aa, and 189–264 aa shortened proteins did not, indicating that the deacetylase activity may be in the 195–245 aa range. Interestingly, although the full amino acid sequence of AhCobQ (1–254 aa) is homologous with Soj protein (A0KQYT), site-determining protein (A0KI19 and A0KK56), and ParA family protein (A0KQG8) in A. hydrophila, the KDAC activity domain (195–245 aa) does not share homology with these proteins, which suggests the different biological roles among them. Moreover, AhCobQ (195–245 aa) also shares homology with several prokaryotic species, such as Aeromonas sp., Aliivibrio sp., Enterovibri sp., Oceanimonas sp., Photobacterium sp., and Vibrio sp., most of which are marine prokaryotes based on a PSI-BLAST program searching, but the search failed to find homologous proteins in eukaryotic cells (Figure 5). Therefore, the 195–245 aa shortened protein of AhCobQ may represent a new domain that exhibits NAD+- and Zn2+-independent deacetylase activity.

Added in the revised version: Lysine deacetylase activity of AhCobQ truncated proteins

Earlier version: Lysine deacetylase activity of AhCobQ truncated proteins. Western blot analysis of the KDAC activity of (A) AhCobQ1–179, AhCobQ179–264, AhCobQ1–264, AhCobQ189–264, AhCobQ199–264, AhCobQ195–264, AhCobQ195–255, and AhCobQ200–255; (B) AhCobQ179–264, AhCobQ189–264, AhCobQ189–255, AhCobQ189–250, AhCobQ189–245, and AhCobQ189–240 truncated proteins treated with Kac-BSA. The first lane represents Kac-BSA without AhCobQ truncated proteins. The upper part of the figure shows the WB results of the Kac level of truncated proteins with Kac-BSA, and the lower part shows the PVDF membrane R350 staining for the loading amount control. (C) Design of the series of truncated proteins and the summary of the WB results in this experiment. Blue indicates the AAA_31 domain (1–179 aa), and red indicates the unknown range (180–264 aa) in AhCobQ. Black and white indicate the positive and negative WB results, respectively, in this study.

Revised version: Western blot analysis of the KDAC activity of (A) AhCobQ1–179, AhCobQ179–264, AhCobQ1–264, AhCobQ189–264, AhCobQ199–264, AhCobQ195–264, AhCobQ195–255, and AhCobQ200–255; (B) AhCobQ179–264, AhCobQ189–264, AhCobQ189–255, AhCobQ189–250, AhCobQ189–245, and AhCobQ189–240 truncated proteins treated with Kac-BSA. The first lane represents Kac-BSA without AhCobQ truncated proteins. The upper part of the figure shows the WB results of the Kac level of truncated proteins with Kac-BSA, and the lower part shows the PVDF membrane R350 staining for the loading amount control. (C) Design of the series of truncated proteins and the summary of the WB results in this experiment. Blue indicates the AAA_31 domain (1–179 aa), and red indicates the unknown range (180–264 aa) in AhCobQ. Black and white indicate the positive and negative WB results, respectively, in this study.

Earlier version: Phylogenetic tree of the AhCobQ KDAC activity domain (195– 245 aa). Different colors indicate different branches, and Aeromonas hydrophila spp. is highlighted in red.

Revised version: Phylogenetic tree of the AhCobQ KDAC activity domain (195–245 aa)

Added in the revised version: Different colors indicate different branches, and Aeromonas hydrophila spp. is highlighted in red.

AhCobQ, AhCobB, and AhAcuC deacetylate different acetylated proteins and sites

Earlier version: We then identified the proteins that were substrates of AhCobQ and analyzed the substrate differences among AhCobQ, AhCobB, and a predicted Zn2+-dependent deacetylase AhAcuC of A. hydrophila that is homologous to type I KDACs in eukaryotes. A quantitative lysine acetylome was generated to compare the differential expression of lysine acetylated peptides among ahcobQ, ahcobB, and ahacuC gene- deficient strains and their WT strain (Figure 6, Figures S8-S110 and Table S2). We only focused on the upregulated lysine acetylated peptides between the mutant and WT strains. When compared to WT, the deletion of ahcobQ, ahcobB, and ahacuC led to increased abundances at 52 Kac sites (corresponding to 47 Kac proteins), 370 Kac sites (291 proteins), and 197 Kac sites (176 proteins), respectively. Moreover, there were 19 increased Kac peptides common between the Δ ahcobQ and Δ ahacuC strains, 10 increased Kac peptides common between the ΔahcobQ and ΔahcobB strains, 69 increased Kac peptides common between the ΔahcobB and ΔahacuC strains, and four peptides common among the ΔahcobQ, ΔahcobB, and ΔahacuC strains (Figure 6A). We also found that there were only 531 Kac peptides directly regulated by these three deacetylases, which is only a small proportion of the total Kac peptides (accounting for 10.28% of all 3,163 Kac peptides identified in the current study), indicating the presence of other unknown KDACs in this bacterium.

Revised version: We then identified the proteins that were substrates of AhCobQ and analyzed the substrate differences among AhCobQ, AhCobB, and a predicted Zn2+-dependent deacetylase AhAcuC of A. hydrophila that is homologous to type I KDACs in eukaryotes. A quantitative lysine acetylome was generated to compare the differential expression of lysine acetylated peptides among ahcobQ, ahcobB, and ahacuC gene-deficient strains and their WT strain (Figure 6, Figures S8-S10 and Table S2). We only focused on the upregulated lysine acetylated peptides between the mutant and WT strains. When compared to WT, the deletion of ahcobQ, ahcobB, and ahacuC led to increased abundances at 52 Kac sites (corresponding to 47 Kac proteins), 370 Kac sites (291 proteins), and 197 Kac sites (176 proteins), respectively. Moreover, there were 19 increased Kac peptides common between the ΔahcobQ and ΔahacuC strains, 10 increased Kac peptides common between the ΔahcobQ and ΔahcobB strains, 69 increased Kac peptides common between the ΔahcobB and ΔahacuC strains, and four peptides common among the ΔahcobQ, ΔahcobB, and ΔahacuC strains (Figure 6A). We also found that there were only 531 Kac peptides directly regulated by these three deacetylases, which is only a small proportion of the total Kac peptides (accounting for 10.28% of all 3,163 Kac peptides identified in the current study), indicating the presence of other unknown KDACs in this bacterium.

Proteomics analysis of upregulated proteins in ΔahcobB, ΔahcobQ, ΔahacuC, and WT strains

(A) Venn diagrams visualized the overlapped upregulated Kac proteins and Kac peptides among the ΔahcobB, ΔahcobQ, ΔahacuC, and WT strains. The number of acetylated proteins and peptides (sites) are shown in the figure. (B) KEGG metabolic pathway enrichment analysis of upregulated Kac proteins and peptides by the three KDACs. Unique and overlapped Kac protein and peptides (sites) are presented in different colors.

A KEGG pathway analysis showed that biosynthesis of amino acids and antibiotics was enriched in the unique upregulated Kac proteins in the ΔahcobQ strain (Figure 6B). Furthermore, five pathways, namely, pyrimidine metabolism, metabolic pathways, purine metabolism, pyruvate metabolism, and RNA polymerase, were enriched in upregulated Kac proteins under AhCobB-regulation. Of these three deacetylase mutants, Δ ahacuC enriched the most metabolic pathways, although Δ ahcobB regulated more Kac sites. Carbon metabolism, biosynthesis of antibiotics, and glycolysis/gluconeogenesis pathways were the top three enrichment pathways in the unique 119 Kac sites regulated by AhAcuC. Phosphotransferase system (PTS), microbial metabolism in diverse environments, and metabolic pathways were enriched in the upregulated Kac proteins common to the ΔahacuC and ΔahcobB mutants. In general, although there are different pathways involved in the activity of different deacetylases, most of the Kac proteins regulated by these KDACs were enzymes and normally had important roles in metabolic pathways.

Validation of the specificity of Kac proteins and sites by deacetylases

Earlier version: We further validated the specificity of Kac proteins and the sites regulated by these deacetylases. A total of 11 target proteins (HrpA, SUN, ENO, ArcA, ADK, IscS, A0KJ75, GyrB, FtsA, RecA, and ICD) were selected for validation due to their significantly upregulated Kac levels in the ΔahcobQ strain, potentially indicating their roles as substrates for AhCobQ. These proteins were first cloned and purified in E. coli. The anti-Kac western blot showed that these purified proteins did not have lysine acetylation signals, which may be because of the heterologous expression in an engineered bacterial strain. This was ideal for excluding signal interference from the background (Figure S11). Next, site-specific lysine acetylation of the target protein substrates was constructed using a two-plasmid-based system of genetically encoded NƐ-acetyllysine, which was purified and incubated with AhCobQ, AhCobB, and AhAcuC for western blot analysis against anti-lysine acetylation antibodies (Figure S12). As shown in Figure 7A, when compared to untreated Kac proteins, the acetylation level of four Kac sites (K816 in HrpA, K103 and K148 in SUN, K195 in ENO) that were uniquely regulated by AhCobQ according to quantitative acetylome analysis was significantly reduced after incubation with AhcobQ but unchanged after incubation with AhCobB or AhAcuC. Moreover, K26 in ArcA-2 and K48 in ADK were reduced in both AhCobB and AhCobQ incubations, and K93 in IscS and K81 in A0KJ75 were decreased with all three deacetylases (Figure 7A). These results were consistent with our quantitative acetylome results, except for K34 in GlmM, which was decreased in AhCobB and AhCobQ treatments in the western blotting assay but upregulated by all three deacetylase mutants according to the proteomics data (Figure S13). Intriguingly, several proteins that contained multiple Kac sites were found to be regulated by different deacetylases. For example, K331 and K449 in GyrB were deacetylated by AhCobQ and AhCobB, respectively (Figure 7B). The same results were found in K320 and K409 in FtsA and in K279 and K306 in RecA as well. These results indicate the different regulatory mechanisms among these KDACs.

Revised version: We further validated the specificity of Kac proteins and the sites regulated by these deacetylases. A total of three target proteins (SUN, ENO, and ArcA-2) were selected for validation due to their significantly upregulated Kac levels in the ΔahcobQ strain, potentially indicating their roles as substrates for AhCobQ. These proteins were first cloned and purified in E. coli. The anti-Kac western blot showed that these purified proteins did not have lysine acetylation signals, which may be because of the heterologous expression in an engineered bacterial strain. This was ideal for excluding signal interference from the background (Figure S11). Next, site-specific lysine acetylation of the target protein substrates was constructed using a two-plasmid-based system of genetically encoded N−-acetyllysine, which was purified and incubated with AhCobQ, AhCobB, and AhAcuC for western blot analysis against anti-lysine acetylation antibodies (Figure S12). As shown in Figure 7, when compared to untreated Kac proteins, the acetylation level of three Kac sites (K103 and K148 in SUN, K195 in ENO) that were uniquely regulated by AhCobQ according to quantitative acetylome analysis was significantly reduced after incubation with AhcobQ but unchanged after incubation with AhCobB or AhAcuC. Moreover, K26 in ArcA-2 was reduced in both AhCobB and AhCobQ incubations. These results were consistent with our quantitative acetylome results.

Added in the revised version: Western blot validation of the site-specific Kac protein substrates regulated by the three KDACs

Removed in the revised version: Western blot validation of the site-specific Kac protein substrates regulated by the three KDACs. (A–C) Kac specific sites were regulated by AhCobQ alone, AhCobQ with AhCobB, and all three KDACs together. The overlapping proteomic features of the target Kac proteins regulated by the three KDACs are shown on the left side of the figure. The detected proteins are in the pink area of the Venn diagram; (D) different Kac sites in one protein substrate were regulated by different KDACs. The predicted 3D structure of the target proteins with different Kac sites is shown on the left side of the figure. The middle section shows the WB results of the deacetylation effect of the three KDACs on site-specific Kac protein substrates, and the right shows the PVDF membrane R350 staining for the loading amount control.

AhCobQ positively regulates the enzymatic activity of isocitrate dehydrogenase (ICD) at K388 site

To further understand the effect of AhCobQ on bacterial biological function, an AhCobQ substrate, A. hydrophila isocitrate dehydrogenase (ICD) was cloned and purified in E. coli (Figure S12). Western blot assay showed that this protein didn’t display lysine acetylation modification in E. coli (Figure 8A). Therefore, site-specific lysine acetylation was constructed at K388 site of ICD for further enzymatic activity determination. As showed in Figure 8B, AhCobQ can significantly deacetylated ICD at K388, while other two deacetylases, AhCobB and AhAcuC didn’t. This result is consistent with the proteomics conclusions. When compared with unKac ICD protein, the Kac at K388 of ICD significantly decreased the product of NADPH with the adding of NADP+ substrate, which indicate the Kac at K388 negatively regulates the enzymatic activity of ICD. Moreover, the adding of AhCobQ significantly increased the enzymatic activity of K388 acetylated ICD and quickly reached the platform period in 10 minutes, while AhCobB and AhAcuC didn’t affect ICD enzymatic activity (Figure 8C and D). These results indicate that AhCobQ positively regulates the enzymatic activity of isocitrate dehydrogenase by deacetylating the ICD protein at K388 site.

Added in the revised version: AhCobQ positively regulates the ICD enzymatic activity through deacetylation of K388 on ICD

Earlier version: AhCobQ positively regulates the ICD enzymatic activity through deacetylation of K388 on ICD. (A) western blot verified the deacetylation effect of AhCobQ on ICD and ICD-K388; (B) western blot verified the deacetylation effect of AhCobQ, AhCobB and AhAcuC on ICD-K388. The PVDF membrane R350 staining for the loading amount control was displayed under the WB results; (C) Enzymatic activities of ICD (yellow), ICD-Kac388(blue), ICD-Kac388 treated with AhCobQ (red), ICD-Kac388 treated with AhCobB (green), ICD-Kac388 treated with AhAcuC (brown); (D) The histogram showed the effect of AhCobQ/AhcobB on ICD or ICD-Kac388 enzymatic activities at 5min. ***P &lt; 0.001.

Revised version: (A) western blot verified the deacetylation effect of AhCobQ on ICD and ICD-K388; (B) western blot verified the deacetylation effect of AhCobQ, AhCobB and AhAcuC on ICD-K388. The PVDF membrane R350 staining for the loading amount control was displayed under the WB results; (C) Enzymatic activities of ICD (yellow), ICD-Kac388(blue), ICD-Kac388 treated with AhCobQ (red), ICD-Kac388 treated with AhCobB (green), ICD-Kac388 treated with AhAcuC (brown); (D) The histogram showed the effect of AhCobQ/AhcobB on ICD or ICD-Kac388 enzymatic activities at 5min. ***P < 0.001.

Discussion

Lysine acetylation has been confirmed to play crucial roles in almost all essential biological functions in both eukaryotic and prokaryotic cells. Normally, reversibility and dynamicity are necessary characteristics of lysine acetylation. The regulatory mechanism of deacetylation has received as much attention as that of acetylation. In different bacterial species, there are hundreds to thousands of protein lysine acetylation modification sites, which are theoretically unlikely to be regulated by one or a few deacetylases, and there must be some unknown deacetylases in bacteria that act on specific protein substrates to maintain the normal physiological function of cells. Therefore, finding new bacterial deacetylases and identifying their modified protein substrates, so as to better understand the molecular mechanism of their regulation, are important in the research on protein acetylation modification regulation.

Here, we reported that CobQ protein in A. hydrophila is an NAD+- and Zn2+-independent deacetylase, and it is very different from currently known deacetylases. We were first interested in the function of AhCobQ based on its annotation as a CobQ/CobB/MinD/ParA family protein in the UniProt database. The deacetylation assay suggested that AhCobQ deacetylated Kac-BSA in the same manner as AhCobB. To our surprise, AhCobQ does not share homology with any prokaryotic sirtuin CobB or currently known eukaryotic deacetylases. We then asked why this was the case since both proteins appeared to belong to the same protein family. Actually, CobB has a conflicting naming phenomenon in the UniProt annotation. There are two distinct CobB annotations in the UniProt database: one is the well-known NAD-dependent protein deacylase, such as NPD_ECOLI in E. coli, and the other is the ATP-dependent a,c-diamide synthase, such as CobB_SINSX in Sinorhizobium sp. that plays an important role in cobalamin biosynthesis (Galperin & Grishin, 2000). In some bacterial species, a,c-diamide synthase is also annotated as CbiA (such as CBIA_SALTY, for example), indicating the controversial assignment of CobB. Interestingly, there is not an a,c-diamide synthase homologous protein in A. hydrophila ATCC 7966, and the CobB annotation in A. hydrophila ATCC 7966 is specific for NAD-dependent protein deacylase. In general, the CobQ protein is annotated as cobyric acid synthase and belongs to a,c-diamide synthase (CobB or CbiA) but not the sirtuin protein (CobB) family. The role of AhCobQ is unknown to date because A. hydrophila may import rather than synthesize cobalamin (vitamin B12) (Seshadri et al, 2006).

The LC MS/MS analysis of synthesized Kac peptides treated by AhCobQ further validated that this protein is a deacetylase and excluded the possibility of it being a hydrolase. Surprisingly, the enzymatic activity of AhCobB was not affected by NAD+, Zn2+, or the sirtuin inhibitor NAM. Given that AhCobQ was reported to be an ATP-dependent cobyric acid synthase (Galperin & Grishin, 2000), we also tested the effect of ATP on deacetylase activity but obtained negative results. Thus, the deacetylase activity of AhCobQ was determined to be NAD+-, Zn2+-, and ATP-independent. Subsequent assays showed that the deacetylase activity of AhCobQ was mediated by an unknown C-terminal domain in the 195–245 aa range that was downstream of the predicted ATPase domain, suggesting a new domain was responsible for deacetylation. However, the intrinsic enzymatic mechanism needs to be further investigated.

Earlier version: We then further asked whether the protein substrates of AhCobQ were the same as the unknown bacterial KDAC homologs in A. hydrophila, AhCobB, and AhAcuC. Although the CobB substrates were identified in E. coli, the AcuC substrates and their comparison with the CobB substrates at a high throughput level had not been determined (Kuhn et al, 2014; Weinert et al, 2017). The characteristics of Kac sites regulated by CobB and AcuC are not yet known. Here, we compared the differential expression of Kac proteins and sites among ahcobQ, ahcobB, and ahacuC mutant and wild-type strains using a quantitative proteomics method. We only focused on the high abundance Kac proteins or sites (fold change > 1.5) when comparing the mutant and WT strains. The results showed that 9.9% of the acetylated proteins (312/3164 Kac protein) were regulated by ahcobB, which is consistent with the proportion in the E.coli cobB mutant and the previous conclusion that CobB is the predominate deacetylase in some bacterial species (Weinert et al, 2013; AbouElfetouh et al, 2015). We also found 47 and 186 significantly enriched Kac proteins in ΔahcobQ and ΔahacuC, respectively.

Revised version: We then further asked whether the protein substrates of AhCobQ were the same as the unknown bacterial KDAC homologs in A. hydrophila, AhCobB, and AhAcuC. Although the CobB substrates were identified in E. coli, the AcuC substrates and their comparison with the CobB substrates at a high throughput level had not been determined (Kuhn et al, 2014; Weinert et al, 2017). The characteristics of Kac sites regulated by CobB and AcuC are not yet known. Here, we compared the differential expression of Kac proteins and sites among ahcobQ, ahcobB, and ahacuC mutant and wild-type strains using a quantitative proteomics method. We only focused on the high abundance Kac proteins or sites (fold change > 1.5) when comparing the mutant and WT strains. The results showed that 9.9% of the acetylated proteins (312/3164 Kac protein) were regulated by ahcobB, which is consistent with the proportion in the E.coli cobB mutant and the previous conclusion that CobB is the predominate deacetylase in some bacterial species (Weinert et al, 2013; AbouElfetouh et al, 2015). We also found 47 and 186 significantly enriched Kac proteins in ΔahcobQ and Δ ahacuC, respectively. Interestingly, there were unique and overlapping Kac substrates among three KDACs, indicating that these KDACs may construct a complicated network for co-regulation of acetylation proteins. Furthermore, although the enriched GO terms and related substrates were different, all the KDACs appeared to be involved in metabolic pathways.

Removed in the revised version: Interestingly, there were unique and overlapping Kac substrates among three KDACs, indicating that these KDACs may construct a complicated network for co-regulation of acetylation proteins. Furthermore, although the enriched GO terms and related substrates were different, all the KDACs appeared to be involved in metabolic pathways.

Removed in the revised version: To further validate our proteomics results, plasmid-based site-specific acetylation was introduced as a target of the recombinant proteins, and the deacetylase activity of the KDACs was determined in vitro. The behavior of the three KDACs on a total of 14 target site-specific acetylated proteins was consistent with our proteomics results, with the results exceeding our expectations. AhCobQ regulated certain proteins at specific Kac sites, which implied that it is involved in specific biological functions. Moreover, it also co-regulated the same Kac sites as did one or all the other KDACs, indicating that the cooperation of these KDACs may play an important role during the maintenance of cellular acetylation homeostasis. Most importantly, our results also showed that AhCobQ deacetylated different Kac sites on the same protein as AhCobB or AhAcuC.

Earlier version: For example, the bacterial topoisomerases DNA gyrase GyrB is an essential enzyme that controls the topological state of DNA during replication (Tari et al, 2013). Previous studies reported the important role of GyrB in bacterial fluoroquinolone resistance (Chauffour et al, 2021; Li et al, 2022). The 3D structure prediction showed that GyrB K331 is located in the ATPase domain that is responsible for supercoiling DNA, whereas K449 is located in the Toprim domain that is required for the interaction with GyrA (Brino et al, 2000). Our current data showed that K331 and K449 were deacetylated by AhCobQ and AhCobB, respectively, indicating the delicate regulatory mechanisms in both functional domains. To our knowledge, this is the first evidence that different KDACs regulate the Kac level at specific sites of the same protein and that they may be involved in differential biological functions in prokaryotic cells. To our knowledge, it is the first evidence that different KADCs regulates the Kac level at specific-site of one protein and that may involve in differential biological functions in prokaryotic cell. We then further wondered the role of AhCobQ in bacterial biological functions. Isocitrate dehydrogenase (ICD) is a key enzyme in the tricarboxylic acid cycle pathway, and also affects bacterial biofilm formation and virulence in several species. Diverse Kac sites of bacterial ICD have been well identified. In E. coli K12, there are at least 15 Kac sites on ICD, and the Kac at K55 and K350 have been determined to play crucial roles on the enzymatic activity by a positively regulatory manner, while Kac at K100, K230 and K235 in a negatively regulatory manner . However, only four Kac sites including K55, K142, K177 and K350 have been proved to be deacetylated by CobB sirtuin obviously. In our quantitative MS data, there are at least ten Kac sites in ICD protein, while only the Kac level of ICD K388 was AhCobQ regulated and others were not altered in KADCs deleted strains. The following assay confirmed that the enzymatic activity of ICD was negatively regulated in a dose- dependent manner by deaectylated ICD K388 site by AhCobQ, whereas not affected by AhCobB and AhAcuC. These results indicate the important role of AhCobQ in bacterial metabolic regulation.

Revised version: To further validate our proteomics results, plasmid-based site-specific acetylation was introduced as a target of the recombinant proteins, and the deacetylase activity of the KDACs was determined in vitro. The behavior of the three KDACs on 4 target site-specific acetylated proteins was consistent with our proteomics results. AhCobQ regulated certain proteins at specific Kac sites, which implied that it is involved in specific biological functions. Moreover, AhCobQ and AhCobB also co-regulated the same Kac sites on ArcA-2 (Arginine deiminase), indicating that the cooperation of these KDACs may play an important role during the maintenance of cellular acetylation homeostasis. Most importantly, our results also showed that AhCobQ deacetylated different Kac sites on the same protein as AhCobB or AhAcuC. For example, the bacterial topoisomerases DNA gyrase GyrB is an essential enzyme that controls the topological state of DNA during replication (Tari et al, 2013). Previous studies reported the important role of GyrB in bacterial fluoroquinolone resistance (Chauffour et al, 2021; Li et al, 2022). The 3D structure prediction showed that GyrB K331 is located in the ATPase domain that is responsible for supercoiling DNA, whereas K449 is located in the Toprim domain that is required for the interaction with GyrA (Brino et al, 2000). Our current data showed that K331 and K449 were deacetylated by AhCobQ and AhCobB, respectively, indicating the delicate regulatory mechanisms in both functional domains. To our knowledge, this is the first evidence that different KDACs regulate the Kac level at specific sites of the same protein and that they may be involved in differential biological functions in prokaryotic cells. To our knowledge, it is the first evidence that different KADCs regulates the Kac level at specific-site of one protein and that may involve in differential biological functions in prokaryotic cell. We then further wondered the role of AhCobQ in bacterial biological functions. Isocitrate dehydrogenase (ICD) is a key enzyme in the tricarboxylic acid cycle pathway, and also affects bacterial biofilm formation and virulence in several species (Venkat S, et al, 2018). Diverse Kac sites of bacterial ICD have been well identified. In E. coli K12, there are at least 15 Kac sites on ICD, and the Kac at K55 and K350 have been determined to play crucial roles on the enzymatic activity by a positively regulatory manner, while Kac at K100, K230 and K235 in a negatively regulatory manner. However, only four Kac sites including K55, K142, K177 and K350 have been proved to be deacetylated by CobB sirtuin obviously. In our quantitative MS data, there are at least ten Kac sites in ICD protein, while only the Kac level of ICD K388 was AhCobQ regulated and others were not altered in KADCs deleted strains. The following assay confirmed that the enzymatic activity of ICD was negatively regulated in a dose-dependent manner by deaectylated ICD K388 site by AhCobQ, whereas not affected by AhCobB and AhAcuC. These results indicate the important role of AhCobQ in bacterial metabolic regulation.

Conclusion

In summary, here, we identified a novel NAD+- and Zn2+-independent deacetylase, AhCobQ, in prokaryotic cells, which does not share homology with current known deacetylases and is unique in prokaryotes. The deacetylase activity of this KDAC is through an unknown domain. Moreover, AhCobQ acetylates at least 47 protein substrates involved in many biological processes, especially metabolic pathways. We provide evidences that AhCobQ negatively regulates the enzymatic activity of bacterial isocitrate dehydrogenase at K388 site. These results indicate that AhCobQ has an important regulatory role in bacterial metabolic pathways. AhCobQ is a novel deacetylase in prokaryotic cells and its biological functions need to be further investigated.

Chemicals and reagents

Ni-NTA agarose beads were purchased from Yeasen Biotechnology; C18 ZipTips from Millipore Corporation; water and acetonitrile from Thermo Fisher Scientific; and formic acid (FA), dithiothreitol, acetone, and iodoacetamide from Sigma Aldrich. Synthesized peptides were purchased from Science Peptide (Shanghai, China) with 95–99% purity. All chemicals and reagents used in this study are listed in Table S3. Critical commercial assay kits used in this study are listed in Table S4.

Methods

Bacterial strains and growth conditions

All strains and vectors involved in this study were preserved or constructed in our laboratory; these are listed in Table S5. A. hydrophila was cultured at 30°C, and E. coli was cultured at 37°C. All bacterial solutions were transferred at 1:100 (v:v; bacterial solution : LB) and cultured continuously for 16 h at 200 rpm, unless otherwise specified.

Generation of gene-deleted strains

According to the whole genome sequence of the A. hydrophila ATCC 7966 strain, upstream and downstream fragments of the target gene of about 500 bp were amplified. The amplified gene fragments were ligated with the digested pRE112 plasmid using the restriction sites SacI and XbaI and then transformed first into E. coli MC1061. The plasmid was collected and transformed into E. coli S17 and then introduced into A. hydrophila by bacterial conjugation. The positive clones were selected on resistant LB plates (100 μg/mL ampicillin and 30 μg/mL chloramphenicol). Finally, the positive mutants from the previous step were spread on LB plates containing 20% sucrose, and a single clone was picked. The positive result was sent to Shanghai Sangon Biological Co., Ltd. for sequencing, and the bacteria were kept at −80°C for later use (Li et al, 2019; Kou et al, 2022).

Physiological phenotype assay

The physiological phenotyping was performed as previously described (Li et al, 2019; Cai et al, 2019). Briefly, for hemolytic and swarming motility assays, bacteria cultured overnight were spotted on 0.7% sheep blood or 0.4% agar plates, respectively, at 30°C for 16 h, and the diameter of the hemolytic or swarming circle was observed and measured in three independent experiments. The bacterial biofilm formation ability was measured with a SpectraMax® i3 multifunctional microplate reader using the crystal violet staining method, and the absorbance at an optical density (OD) of 595 nm wavelength was measured, as previously described (Li et al, 2019). All these experiments were independently repeated three times, and the one-way ANOVA method was used for statistical analysis.

BSA acetylation in vitro

Briefly, 1 mg of BSA powder was dissolved in 1 mL of ddH2O containing 0.1 M Na2CO3 and 0.04 g of acetic anhydride and then incubated at 37°C to react for 3 h to produce Kac-BSA. Finally, 20 µL of 1 M Tris (pH 8.0) was added to stop the reaction (Guan et al, 2010).

Protein deacetylation assay in vitro

Briefly, 1 µg of protein or peptides were mixed with 0.2 µg of purified KDAC protein in basic buffer (150 mM Tris-HCl, pH 8.0, and 10% glycerol) with 10 mM ZnCl2, 0.5 mM ATP, 0.25 mM NAD+, or 0.5 mM NAM, as appropriate, and then incubated at 37°C for 3 h (Li et al., 2024).

Quantitative acetylome analysis

The quantitative acetylome was performed as previously described (Zhang et al, 2022). Briefly, the protein samples from bacteria cultured overnight were collected by centrifugation and then sonicated in an ice bath. A 10 mg protein sample was reduced by 10 mM dithiothreitol (DTT), alkylated by 20 mM iodoacetamide (IAA), and then digested to peptides by trypsin at a 1:50 ratio at 37°C for 16 h. The lysine acetylation peptides were enriched by an immuno-affinity enrichment method using an agarose-conjugated anti-lysine-acetylated antibody (Hangzhou Jingjie, PTM 101).

LC MS/MS

For BSA acetylation identification, the digested peptides were analyzed on an Orbitrap Q Exactive HF mass spectrometer (Thermo Fisher Scientific) equipped with an EASY-nLC 1200 system (Thermo Fisher Scientific). The peptides were first loaded into a C18 trap column (1.9 μm, 100 μm, 20 mm) and separated on an EASY-nLC 1200 UPLC system. The precursor ions were selected for fragmentation in the higher energy collision-activated dissociation (HCD) cell at a normalized collision energy of 27% and then detected using the Orbitrap analyzer at a resolution of 15,000 at m/z 200.

The synthetic peptides YNGVFQECCQAEDKGACLLPK (STD) and YNGVFQECCQA EDKacGACLLPK (STD*Kac) were analyzed using an Orbitrap Fusion mass spectrometer (Thermo Fisher Scientific) equipped with an EASY-nLC 1200 system (Thermo Fisher Scientific). For the quantitative acetylome, the enriched peptides were analyzed using an Orbitrap Exploris 480 mass spectrometer equipped with a high-field asymmetric waveform ion mobility (FAIMS) Pro spectrometer and an EASY-nLC 1200 system (Thermo Fisher Scientific).

The data obtained were compared and annotated with the BSA (UniProt ID P02769) or A. hydrophila ATCC 7966 database from the UniProt database using MaxQuant v1.6.2.10 software. The mass error of the search was set to 10 ppm, the main search was set to 5 ppm, the fragment ion was set to 0.02 Da, and the variable modification was acetylation modification. The maximum false discovery rate (FDR) threshold for proteins, peptides, and modification sites was set at 1%. In addition, the localization probability threshold was specified as >0.75, and the modification site score threshold was set to 40. The precursor area for quantification of synthetic peptides YNGVFQECCQAEDKGACLLPK (STD) and YNGVFQECCQAEDKacGACLLPK (STD*Kac) was manually exacted with Thermo Xcalibur (v.3.0.63). When performing a label-free quantitative (LFQ) comparison, in addition to the above settings, label-free quantification was also set to LFQ, and the intensity signal ratio of the corresponding modified peptides in the identification results (ratio) was >1.5. P<0.05 (Student’s t test) among the three biological replicates in each group was considered a significant difference.

Western blot

The protein samples (0.2 µg) were separated on a 12% sodium dodecyl-sulfate polyacrylamide gel electrophoresis (SDS-PAGE) gel and then transferred to polyvinylidene fluoride (PVDF) membranes using a semi-dry transfer system (Bio-Rad, USA) at 25 V for 20 min. The membranes were blocked with Tris-buffered saline with 0.1% Tween 20 (TBST) containing 5% skim milk for 1 h and incubated with the diluted primary antibody (PTM-101 1:1,000) for 5 h at room temperature. After washing three times with TBST, the membranes were incubated with the secondary antibody, horseradish peroxidase-conjugated goat anti-mouse IgG (1:5,000), at room temperature for 1 h. Finally, the membranes were covered with Clarity western ECL (Bio-Rad) chromogenic solution and exposed with a ChemiDoc MP (Bio-Rad, Hercules, CA, USA). The exposed PVDF membranes were stained with Coomassie R-350 to verify whether the loading amounts were consistent (Peng et al, 2019).

Dot blot

Removed in the revised version: 10 μg protein was diluted and spotted on the PVDF membrane. Once the samples had air-dried, the membranes underwent a blocking process with Tris-buffered saline with Tween 20 (TBST) containing 5% BSA for a duration of 3 to 5 hours. Following this, they were incubated with the PTM-101 antibody at a 1:1000 dilution for 5 hours.

Earlier version: Afterward, the membranes were washed three times in TBST, then incubated with a secondary antibody, goat anti-mouse HRP-IgG, at a 1:5000 dilution and containing 3% BSA, for one hour at room temperature. This was followed by five washes in TBST, each lasting 5 minutes. To conclude the process, the membrane was treated with Clarity Western ECL Substrate (Bio-Rad) chromogenic solution, exposed using a ChemiDoc MP imaging system (Bio-Rad, Hercules, CA, USA), and finally stained with Coomassie R-350 to ensure consistent sample loading.

Revised version: 10 μg protein was diluted and spotted on the PVDF membrane. Once the samples had air-dried, the membranes underwent a blocking process with Tris-buffered saline with Tween 20 (TBST) containing 5% BSA for a duration of 3 to 5 hours. Following this, they were incubated with the PTM-101 antibody at a 1:1000 dilution for 5 hours. Afterward, the membranes were washed three times in TBST, then incubated with a secondary antibody, goat anti-mouse HRP-IgG, at a 1:5000 dilution and containing 3% BSA, for one hour at room temperature. This was followed by five washes in TBST, each lasting 5 minutes. To conclude the process, the membrane was treated with Clarity Western ECL Substrate (Bio-Rad) chromogenic solution, exposed using a ChemiDoc MP imaging system (Bio-Rad, Hercules, CA, USA), and finally stained with Coomassie R-350 to ensure consistent sample loading.

Recombinant protein production and purification

The target genes were amplified and ligated to the digested pET-21b/32a plasmid using the restriction sites EcoRI and HindIII. Then, the recombinant plasmid was transformed into competent E. coli BL21 (DE3), and the positive clones were selected from an agar plate with 100 μg/mL ampicillin. For the construction of site-directed mutant strains, the target gene was mutated using a Fast Mutagenesis System kit according to the manufacturer’s protocol (Transgen Biotech Co., Beijing, China). All primer pairs used in this study are listed in Table S6.

The successfully constructed recombinant strains were transferred to 200 mL of LB medium containing 100 μg/mL ampicillin, and then, 200 μL of 1 M isopropyl β-D-1-thiogalactopyranoside (IPTG) was added until the OD at 600 nm reached about 0.6. After washing three times with phosphate-buffered saline (PBS) and dissolving in 10 mL of a solution (5 mM imidazole, 0.5 mM NaCl, 0.05M Tris-HCl pH=8.0), the bacterial cells were ultrasonically disrupted for 20 min on ice. The supernatant was collected by centrifuging at 10,000 rpm at 4°C for 15 min, and the recombinant proteins were purified using an Ni+-NTA column as previously described (Jiang et al, 2022).

Expression and purification of site-specific Kac proteins

We used two genetic code expansion systems in E. coli BL21 to express site-specific acetylated proteins (Vidali G et al, 1968; Bryson et al, 2017). First, the target gene was inserted into the pET-21b plasmid as described above, and the lysine codon at the target position of the gene was substituted with an amber stop (TAG) codon by site-directed mutagenesis. Then, equal amounts of pTECH-MbAcK3RS (IPYE) and recombinant pET-21b plasmid were co-transfected into competent E. coli BL21 (DE3) on agar plates containing 100 μg/mL ampicillin and 30 μg/mL chloramphenicol. After culturing for 5 h at 37°C and 200 rpm, the successful PCR products were confirmed with DNA sequencing.

The successfully site-specific acetylated mutants were cultured overnight in 5 mL of LB medium and then transferred to 200 mL of fresh LB medium containing 100 μg/mL ampicillin and 30 μg/mL chloramphenicol and incubated at 200 rpm and 37°C. To express the acetyl proteins, 2 mM N-acetyllysine and 20 mM NAM were added to the medium at 37°C and vortexed at 200 rpm until the OD at 600 nm reached about 0.8. Then, 200 μL of 1M IPTG was added to the medium, and the cells were cultured at 16°C and vortexed at 200 rpm for 18 h to express the proteins. The procedure of site-specific Kac protein purification was the same as the method described above.

Bioinformatics and data analysis

Earlier version: Amino acid sequences were obtained from the UniProt database (https://www.uniprot.org/uniprot). The pfam database (https://www.pfam.org/) was used to predict protein domains, and the online software iTOL (https://itol.embl.de/) was used to generate graphs. The amino acid sequence homology analysis was carried out using position-specific iterated BLAST (PSI-BLAST) software after two iterations (https://blast.ncbi.nlm.nih.gov/Blast.cgi). MEGA v7.0.26 software was used to construct the evolutionary tree that was visualized by online Chiplot software (https://www.chiplot.online/). The protein 3D structure was predicted by AlphaFold in the UniProt database and was visualized by PyMMOL v.2.5.2. The online software STRING v.11.5 (http://string-db.org/) was used to perform Kyoto Encyclopedia of Genes and Genomes (KEGG) metabolic pathway enrichment analysis for upregulated proteins, and these proteins were finally visualized using Cytoscape v.3.9.0. Except where otherwise stated, each experiment in this study was independently repeated three times, and the one-way ANOVA method was used for statistical analysis.

Revised version: Amino acid sequences were obtained from the UniProt database (https://www.uniprot.org/uniprot). The pfam database (https://www.pfam.org/) was used to predict protein domains, and the online software iTOL (https://itol.embl.de/) was used to generate graphs. The amino acid sequence homology analysis was carried out using position-specific iterated BLAST (PSI-BLAST) software after two iterations (https://blast.ncbi.nlm.nih.gov/Blast.cgi). MEGA v7.0.26 software was used to construct the evolutionary tree that was visualized by online Chiplot software (https://www.chiplot.online/). The online software STRING v.11.5 (http://string-db.org/) was used to perform Kyoto Encyclopedia of Genes and Genomes (KEGG) metabolic pathway enrichment analysis for upregulated proteins, and these proteins were finally visualized using Cytoscape v.3.9.0. Except where otherwise stated, each experiment in this study was independently repeated three times, and the one-way ANOVA method was used for statistical analysis.

Data availability

The raw MS data were deposited in the public ProteomeXchange Consortium via the PRIDE partner repository with the dataset identifier PXD038735 or IPX0005366000 (iProx database).

Isocitrate dehydrogenase (ICD) enzymatic activity assay

The enzymatic activity of ICD was detected by measuring the amount of NADPH generated by NADP substrate. 10 μg purified ICD or site-specifically Kac ICD protein was mixed with or without 3 μg purified AhCobQ in 180 μL solution buffer (100 mM Tris-HCl, pH 8.0, 0.5 mM NADP, 3 mM MgCl2, 10% glycerol, 1.5 mM isocitrate), and the absorbance of NADPH at OD340nm was detected by SpectraMax® i3(Molecular Devices) once a minute. Simultaneously, serial concentrations of NADPH (0-0.8mM) was incubated with ICD for standard curve determination.

Removed in the revised version: Supporting information